# Assignment 2

## Deadline

Wednesday, October 15, 23:59.

# Instructions

Write a program that calculates how long it would take to sail around the Earth along the Tropic of Cancer (*obratník Raka*) with a ship moving at a given constant speed.
- Use the `input` function to read the ship’s speed (a real number in kilometers per hour). Use the input function **exactly once** in your program.
- Assume that the Earth is a sphere with a radius R of **6378 km**, and that the ship sails along the parallel of **23.5°** north latitude (the Tropic of Cancer). Radius of this parallel is given by the formula r = R cos(φ), 
where φ is the latitude in radians and R is the radius of the Earth. Make sure to use a sufficiently precise value of π.
- Compute the total travel time in hours, round it to the nearest integer, and print the result using the `print` function in the format `X d Y hrs`, where `X` is the number of days and `Y` is the number of hours.
- For example, for a speed of **53 km/h**, the trip takes approximately **693.4 hours**, so the program should print `28 d 21 hrs`.
- Pay attention to spaces between numbers and the units `d` and `hrs`! Due to automatic checking, **do not print anything else** besides these two numbers and the units.

*Final tip:* don’t forget to save your final solution and click the **Validate** button before submitting.


```python
# this is a read-only cell - do not edit, delete, or move!
import zprotests
```

<!-- grade_id: cell-28dbe96b036581b3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
from math import cos, pi

## length
speed = float(input())
ourRadius = 6378 * cos(23.5 * pi / 180)
circumference = 2 * pi * ourRadius

## time
hTime = circumference / speed
dayTime = int(hTime // 24)
hTime = round(hTime % 24)

## output
print(f"{dayTime} d {hTime} hrs")
```

Output:
```
2
765 d 15 hrs
```

<!-- grade_id: cell-700f2e2b81a58380, points: 25, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!

# This cell validates the example values given in the text of the assignment.
# If something fails and you think that the result is correct, check that you
# use the input and print functions as required by the assignment.

with zprotests.mock_input(inputs=[53]):
    %rerun -g 'input('

assert zprotests.stdout == "28 d 21 hrs\n"
```

<!-- grade_id: cell-9381416308ab808f, points: 25, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!

```

<!-- grade_id: cell-372a56fbe26c4579, points: 20, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!

```

<!-- grade_id: cell-ba081c9478a9d5d7, points: 15, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!

```

<!-- grade_id: cell-d007813640d92e69, points: 15, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!

```