# Assignment 5

## Deadline

Wednesday, November 5, 23:59.

## Task

The [Koch snowflake](https://en.wikipedia.org/wiki/Koch_snowflake) is a fractal curve formed by starting with an equilateral triangle and recursively modifying each straight line segment. At each iteration, every segment is replaced by four segments, each 1/3 the length of the original, forming an equilateral "bump":

![The first seven iterations in animation](https://upload.wikimedia.org/wikipedia/commons/f/fd/Von_Koch_curve.gif)

The **perimeter length** of the Koch snowflake after $n$ iterations follows a recursive pattern:   

- The initial equilateral triangle (base case $n = 1$) has side length $s$, so the perimeter length is $3s$.
- In each iteration ($n > 1$), each line segment of length $s$ is replaced by 4 segments of length $s/3$.

Your task is to write a **recursive function** that calculates the perimeter length of the Koch snowflake with initial side length $s$ after $n$ iterations.

**Comments and hints:**

- The function is defined for integer $n \ge 1$ and real $s \ge 0$. If some of these preconditions does not hold, the function should return `None`.
- If the input parameters are correct, the function should return a `float` or `int` number.
- **Do not use the `print` and `input` functions in this assignment.**
- Do not use any imports in this assignment.

<!-- grade_id: cell-28dbe96b036581b3, points: 0, grade: false, solution: true -->
```python
def koch_length(n, s):
     if type(n) != int or (type(s) != int and type(s) != float) or n < 1 or s < 0:
          return None
     if n == 1:
          return 3 * s
     else:
          return 4/3 * koch_length(n-1, s)
```

Output:
```
142.0618130705007
```

The following cells contain public tests that you can use for basic validation of your solution. **Click on the validate button before submitting!**

```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
import ipytest
ipytest.autoconfig(raise_on_error=True, addopts=["-p", "no:cacheprovider"])  # type: ignore
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[3], line 2
      1 # This is a read-only cell used for automatic validation - do not edit, delete, or move!
----> 2 import ipytest
      3 ipytest.autoconfig(raise_on_error=True, addopts=["-p", "no:cacheprovider"])  # type: ignore

ModuleNotFoundError: No module named 'ipytest'
```

<!-- grade_id: cell-5443d33a7c80dc12, points: 10, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
assert koch_length(1, 0) == 0
assert koch_length(2, 0) == 0
assert koch_length(3, 0) == 0
```

<!-- grade_id: cell-799821c89e9f18c0, points: 10, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
assert koch_length(1, 1) == 3
assert koch_length(1, 10) == 30
assert koch_length(1, 100) == 300
```

<!-- grade_id: cell-b4eeb6797b63d39d, points: 10, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
assert koch_length(2, 1) == 4
assert koch_length(2, 2) == 8
assert koch_length(2, 3) == 12

assert koch_length(3, 1) == 16/3
assert koch_length(3, 2) == 32/3
assert koch_length(3, 3) == 16
```

<!-- grade_id: cell-74cc1e48d19a1fc9, points: 10, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
assert koch_length(2, -1) is None
assert koch_length(-2, -1) is None
assert koch_length(-2, 1) is None

assert koch_length(0, 0) is None
assert koch_length(0, 1) is None
assert koch_length(0, 2) is None

assert koch_length(2.1, 1) is None
```

<!-- grade_id: cell-b163e01d7807fc11, points: 20, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
from test_recursion import test_recursion
assert test_recursion(lambda: koch_length(5, 3))
```

The following cells contain hidden tests that are evaluated during automatic grading after your submission. The tests check that your functions are sufficiently general.

<!-- grade_id: cell-08758ff656ac355b, points: 30, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
```

<!-- grade_id: cell-2acbabe27592dbd2, points: 10, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
```