# Assignment 6

## Deadline

Wednesday, November 12, 23:59.

## Task

You have been hired by a meteorological agency to analyze historical weather data.
The data consists of daily temperature records stored as a list of tuples, where each tuple contains the **day identifier** and **temperature in Celsius**.
Your task is to write functions that analyze this data to identify meaningful patterns. 

### Sample data format

```python
weather_data = [
    ("Monday", 22.5),
    ("Tuesday", 25.3),
    ("Wednesday", 27.1),
    ("Thursday", 24.8),
    ("Friday", 26.0),
    ("Saturday", 28.4),
    ("Sunday", 23.7)
]
```

### Requirements

- **Function 1:** `average_temp(data)`
    - Calculate the average temperature for the week
    - Example: `average_temp(weather_data) → 25.6`

- **Function 2:** `hot_days(data, threshold=25.0)`
    - Return a new list of tuples containing only days with temperatures above threshold
    - Example: `hot_days(weather_data, 26.0) → [("Wednesday", 27.1), ("Friday", 26.0), ("Saturday", 28.4)]`

- **Function 3:** `temp_range(data)`
    - Calculate the temperature range (difference between highest and lowest temperatures)
    - Example: `temp_range(weather_data) → 5.9`

- **Function 4:** `temperature_summary(data, hot_threshold=25.0)`
    - Return a tuple containing:
        1. Average temperature (using `average_temp`)
        2. List of hot days (using `hot_days`)
        3. Temperature range (using `temp_range`)
    - Example: `temperature_summary(weather_data, 26.0) → (25.6, [("Wednesday", 27.1), ("Friday", 26.0), ("Saturday", 28.4)], 5.9)`

### Implementation guidelines 

- Do not use external libraries (only standard Python)
- Do not use global variables
- Do not modify the original data
- Edge cases: Handle empty lists (return `None` or an empty container as appropriate, check the tests!)
- Write a docstring for each function explaining its purpose and usage

<!-- grade_id: cell-28dbe96b036581b3, points: 0, grade: false, solution: true -->
```python
def average_temp(data):
     """
     Calculates the average temperature from weather data

     Input: List of tuples in format ("day", temperature)

     Output: Average temperature
     """

     if not data:
          return None

     total = 0
     for i in range(len(data)):
          total += data[i][1]
     return total / len(data)

def hot_days(data, threshold=25.0):
     """
     Writes the days at least as hot as the threshold

     Input: List of tuples in format ("day", temperature) & the temperature threshold

     Output: Array of the same format as input
     """

     if type(data) != list:
          return None
     hotDays = []

     for i in range(len(data)):
          if threshold <= data[i][1]:
               hotDays.append(data[i])
     return hotDays

def temp_range(data):
     """
     Calculates the temperature range (difference between highest and lowest temperatures)

     Input: List of tuples in format ("day", temperature)

     Output: Difference between the hottest and the coldest day (number)
     """

     if not data:
          return None
     
     lowest = data[0][1]
     highest = data[0][1]

     for i in range(len(data)):
          if lowest > data[i][1]:
               lowest = data[i][1]
          if highest < data[i][1]:
               highest = data[i][1]
     return highest - lowest

def temperature_summary(data, hot_threshold=25.0):
     """
     Summarizes outputs from average_temp(), hot_days() & temp_range()

     Input: List of tuples in format ("day", temperature) & the temperature threshold

     Output: Tuple with respective outputs from average_temp(), hot_days() & temp_range()
     """

     return (average_temp(data), hot_days(data, hot_threshold), temp_range(data))
```

The following cells contain public tests that you can use for basic validation of your solution. **Click on the validate button before submitting!**

```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
import ipytest
ipytest.autoconfig(raise_on_error=True, addopts=["-p", "no:cacheprovider"])  # type: ignore
```

<!-- grade_id: cell-5443d33a7c80dc12, points: 10, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
assert average_temp([("Mon", 22.5), ("Tue", 25.3), ("Wed", 27.1)]) == (22.5 + 25.3 + 27.1) / 3
assert average_temp([("Monday", 25.0)]) == 25.0
assert average_temp([]) is None
```

<!-- grade_id: cell-799821c89e9f18c0, points: 10, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
assert hot_days([("Mon", 22.5), ("Tue", 25.3), ("Wed", 27.1)], 26.0) == [("Wed", 27.1)]
assert hot_days([("Mon", 22.5), ("Tue", 25.3), ("Wed", 27.1)], 20.0) == [("Mon", 22.5), ("Tue", 25.3), ("Wed", 27.1)]
assert hot_days([("Mon", 22.5), ("Tue", 25.3), ("Wed", 27.1)], 30.0) == []
assert hot_days([]) == []
```

<!-- grade_id: cell-b4eeb6797b63d39d, points: 10, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
assert temp_range([("Mon", 22.5), ("Tue", 25.3), ("Wed", 27.1)]) == (27.1 - 22.5)
assert temp_range([("Monday", 25.0)]) == 0.0
assert temp_range([]) is None
```

<!-- grade_id: cell-74cc1e48d19a1fc9, points: 10, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
assert temperature_summary([("Mon", 22.5), ("Tue", 25.0), ("Wed", 27.5)]) == (25.0, [("Tue", 25.0), ("Wed", 27.5)], 5.0)
assert temperature_summary([("Monday", 25.0)]) == (25.0, [("Monday", 25.0)], 0.0)

assert temperature_summary([("Mon", 22.5), ("Tue", 25.0), ("Wed", 27.5)], 26.0) == (25.0, [("Wed", 27.5)], 5.0)
assert temperature_summary([("Monday", 25.0)], 26.0) == (25.0, [], 0.0)

assert temperature_summary([]) == (None, [], None)
```

<!-- grade_id: cell-b163e01d7807fc11, points: 10, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
assert average_temp.__doc__, "average_temp function must have a docstring"
assert hot_days.__doc__, "hot_days function must have a docstring"
assert temp_range.__doc__, "temp_range function must have a docstring"
assert temperature_summary.__doc__, "temperature_summary function must have a docstring"
```

The following cells contain hidden tests that are evaluated during automatic grading after your submission. The tests check that your functions are sufficiently general.

<!-- grade_id: cell-08758ff656ac355b, points: 50, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
```