# Assignment 9

## Deadline

Wednesday, December 3, 23:59.

## Task: Parrot Speech Analysis

You are given a dictionary mapping parrot names (`str`) to their speeches (`str`).  
Example:
```python
parrot_speeches = {
    "Charlie": "Charlie loves crackers.  Crackers are tasty, and Charlie wants more crackers. Hello, Charlie loves tasty crackers!",
    "Max": "Hello Max, hello! Bad Max, Max, Max. Max? Hello!",
    "Milo": "Milo loves shiny things.  Shiny, shiny things make Milo happy. Milo loves sunny days and pretty things.",
    "Polly": "Hello! Hello! Polly, sweet Polly. Tasty crackers, tasty crackers. Hello, Polly wants a sweet cracker! Cracker!",
}
```
### Word-processing rules
In all tasks, the speeches should be converted into lists of words following the rules:
- A __word__ is a __non-empty__ sequence of letters (`a–z`, `A–Z`) separated by whitespace.
- __Punctuation:__ remove characters such as `.,!?'"` and all other non-letter characters  (`polly's` -> `pollys`).
- __Case-insensitive:__ convert all words to lowercase (`Cracker` -> `cracker`).
- __Different grammatical forms__ count as different words (`cracker` ≠ `crackers`).


### Functions to implement
1. `speech_to_words(speech)`
2. `parrot_vocabularies(speeches)`
3. `who_knows_word(speeches, word)`
4. `pairs_with_shared_vocabulary(speeches, min_shared)`

#### 1. `speech_to_words(speech)`  (helper function)
- This function must correctly convert the speech into a list of words satisfying the word-processing rules listed above:
   1. Remove punctuation characters and any other non-letter characters.
   2. Convert all words to lowercase.
   3. Return words in the original order.
- Example: `speech_to_words("Hello, Polly wants a cracker. Cracker!")` -> `["hello", "polly", "wants", "a", "cracker", "cracker"]`

#### 2. `parrot_vocabularies(speeches)`
- Find parrots' _vocabularies_. For this task, the _vocabulary_ of a parrot is the set of distinct words it uses in its speech (use `speech_to_words` to preprocess the speeches).
- Return a sorted list of tuples `(name, vocabulary)`
- Where:
   - `name` is a parrot name,
   - `vocabulary` is a set of distinct words used by the parrot.
- Sorting rules: 
    1. Descending by size of the `vocabulary`
    2. Ascending alphabetically by `name` (if sizes match).
- Example output:
`
[('Milo', {'and', 'happy', 'sunny', 'loves', 'days', 'things', 'make', 'milo', 'shiny', 'pretty'}),
  ('Charlie', {'crackers', 'and', 'charlie', 'hello', 'loves', 'are', 'wants', 'tasty', 'more'}),
  ('Polly', {'crackers', 'cracker', 'hello', 'a', 'polly', 'wants', 'tasty', 'sweet'}),
  ('Max', {'hello', 'bad', 'max'})]
 `

#### 3. `who_knows_word(speeches, word)`
- Find parrots that used the given `word` at least once (use `speech_to_words` to preprocess the speeches).
- Return a sorted list of tuples `(name, count)`
- Where:
   - `name` is a parrot name,
   - `count` is the number of occurences of the `word` of the parrot`s speech.
- Include only parrots with at least one occurrence of the word.
- Sorting rules: 
   - 1. Descending by `count`,
   - 2. Ascending alphabetically by `name` if counts match.
- Example: `who_knows_word(parrot_speeches, "hello")` -> `[('Max', 3), ('Polly', 3), ('Charlie', 1)]`

#### 4. `pairs_with_shared_vocabulary(speeches, min_shared)`
- Find all pairs of parrots that share at least `min_shared` common words in their vocabularies (use `speech_to_words` to preprocess the speeches).
- Return a dictionary where each key is a pair of parrot names
`(name1, name2)` and the corresponding value is the set of shared words.
- Rules:
    - Include only pairs with at least `min_shared` shared words.
    - Always ensure `name1 < name2` alphabetically. (This guarantees that each pair appears exactly once and avoids duplicates)
- Example: `pairs_with_shared_vocabulary(parrot_speeches, 2)` -> `{('Charlie', 'Milo'): {'and', 'loves'}, ('Charlie', 'Polly'): {'tasty', 'hello', 'wants', 'crackers'}}`

## Additional notes
- Do not use external libraries (only standard Python)
- Do not use global variables
- Do not modify the original data (i.e., `speeches`)
- You may assume that `speeches` is always a dictionary where both keys and values are of type `str`.
- Do not use the `input` or `sys.exit` functions in your solution.
- Test your code carefully before submitting.


<!-- grade_id: cell-6cb4486c1393db6c, points: 0, grade: false, solution: true -->
```python
#

#! Import
from re import fullmatch, sub
from itertools import combinations

#! Convert Sentence Into Array Of Words
def speech_to_words(speech):
     """
     Convert a speech into a list of words.

     1. Remove punctuation and any other non-letter characters.
     2. Convert all words to lowercase.
     3. Return the list of words in the order they appear.
     """
     lowercase = sub(r"\s+", " ", speech.lower()) + " " # add a space so even the last word is add to wordArray
     letterArray = []
     wordArray = []

     for char in lowercase:
          if fullmatch(r"[a-z]", char):
               letterArray.append(char)
          elif fullmatch(" ", char):
               if "".join(letterArray) != "":
                    wordArray.append("".join(letterArray))
               letterArray = []
     return wordArray

#! Transfer Dictionary Into Array Of (Names, {Vocabulary})
def parrot_vocabularies(speeches):
     """
     Return a list of (name, vocabulary_set) for each parrot.

     Vocabulary = set of distinct words in the parrot's speech.

     Sorting:
       1. Descending by vocabulary size.
       2. Ascending alphabetically by name when equal.
     """
     parrotVocabArray = []

     #* Create Array
     for papagei in speeches:
          wordlist = speech_to_words(speeches[papagei])
          parrotVocabArray.append((papagei, set(wordlist)))
     
     #* Sort Array
     parrotVocabArray.sort(key=lambda x: x[0]) # alphabetically
     parrotVocabArray.sort(key=lambda x:len(x[1]), reverse=True) # by richness of vocabulary

     return parrotVocabArray

#! Count "Word" Uses Per Papagei    
def who_knows_word(speeches, word):
     """
     Find parrots that used the given `word` at least once.

     Return a list of (name, count), where:
          - name is a parrot name,
          - count is the number of occurrences of `word` in that parrot's speech.

     Only include parrots with count >= 1.

     Sorting:
       1. Descending by count.
       2. Ascending by parrot name when counts are equal.
     """
     array = []
     word = word.lower()

     #* Find number of "word" instances
     for papagei in speeches:
          wordlist = speech_to_words(speeches[papagei])
          array.append((papagei, wordlist.count(word)))

     #* Filter & sort
     filteredArray = list(filter(lambda x: x[1] != 0, array)) # delete those with 0 occurances
     filteredArray.sort(key=lambda x: x[0]) # alphabetically
     filteredArray.sort(key=lambda x: x[1], reverse=True) # by richness of vocabulary

     return filteredArray

#! Find Shared Vocab Pairs
def pairs_with_shared_vocabulary(speeches, min_shared):
     """
     Find all pairs of parrots that share at least `min_shared` words.

     Return a dictionary of the form: { (name1, name2): shared_words_set, ... }
     Where:
          - name1 and name2 are two different parrot names,
          - name1 < name2 alphabetically (to avoid duplicate pairs),
          - shared_words_set is the set of words used by both parrots.
     A pair is included only if the size of shared_words_set is >= min_shared.
     """
     grandObject = {}

     #* Find all used words
     allWordsUsed = []
     for papagei in speeches:
          currentWordlist = speech_to_words(speeches[papagei])
          allWordsUsed += currentWordlist
     allWordsUsed = list(set(allWordsUsed))

     #* List papageis
     papageis = []
     for papagei in speeches:
          papageis.append(papagei)

     #* Find all papagei combinations
     papageiCombo = list(combinations(sorted(papageis), 2))

     #* Find shared words
     for word in allWordsUsed:
          whoKnowsArrayWithCounts = who_knows_word(speeches, word)

          if len(whoKnowsArrayWithCounts) >= 2:
               whoKnowsArray = []

               for tup in whoKnowsArrayWithCounts:
                    whoKnowsArray.append(tup[0])
               knowsThisWordCombo = list(combinations(sorted(whoKnowsArray), 2))

               for papageiDuo in papageiCombo:
                    for knowsDuo in knowsThisWordCombo:
                         if papageiDuo == knowsDuo:
                              if papageiDuo not in grandObject:
                                   grandObject[papageiDuo] = {word,}
                              else:
                                   grandObject[papageiDuo].add(word)
     
     #* Filter pairs with fewer than min_shared words
     markedForDeletion = []
     for pair in grandObject:
          if len(grandObject[pair]) < min_shared:
               markedForDeletion.append(pair)

     for pair in markedForDeletion:
          del grandObject[pair]
     return grandObject
```

# Tests
The following cells contain public tests that you can use for basic validation of your solution. **Click on the validate button before submitting!**

<!-- grade_id: cell-3eb19aec20bb3f75, points: 0, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
# Functions to run tests
from pprint import pprint

def check_speech_to_words(speech, solution):
    """Check the speech_to_words(speech) function."""
    result = speech_to_words(speech)
    try:
        assert isinstance(result, list), "Result should be a list"
        for w in result:
            assert isinstance(w, str), "Each item in the result should be a string (word)"
        assert result == solution, "The result is not correct"
    except AssertionError as e:
        print(f"Test Failed (speech_to_words): {e}")
        print("Test case:")
        pprint(speech)
        print("Your solution:")
        pprint(result)
        print("Correct solution:")
        pprint(solution)
        raise
    print("test OK (speech_to_words)")


def check_parrot_vocabularies(speeches, solution):
    """ Check the parrot_vocabularies(speeches) function. """
    speeches_orig = dict(speeches.items())
    result = parrot_vocabularies(speeches)

    try:
        assert isinstance(result, list), "Result must be a list."
        for item in result:
            assert isinstance(item, tuple) and len(item) == 2, \
                "Each item must be a (name, vocabulary_set) tuple."
            assert isinstance(item[0], str), "Name must be a string."
            assert isinstance(item[1], set), "Vocabulary must be a set."
            for w in item[1]:
                assert isinstance(w, str), "Vocabulary items must be strings."
        assert sorted(speeches_orig.items()) == sorted(speeches.items()), \
            "Your function makes changes to speeches."
        assert len(result) == len(solution), \
            "Incorrect number of items in result."
        assert sorted(name for name, _ in result) == \
               sorted(name for name, _ in solution), \
               "Incorrect parrot names in result."
        assert sorted(v for _, v in result) == \
               sorted(v for _, v in solution), \
               "Incorrect vocabulary sets."
        assert result == solution, "Result is not sorted correctly."

    except AssertionError as e:
        print("Test FAILED:", e)
        print("Test case speeches:")
        pprint(speeches)
        print("Your result:")
        pprint(result)
        print("Expected result:")
        pprint(solution)
        raise

    print("test OK")

def check_who_knows_word(speeches, word, solution):
    """Check the who_knows_word(speeches, word) function."""
    speeches_orig = dict(speeches.items())
    result = who_knows_word(speeches, word)

    try:
        assert sorted(speeches_orig.items()) == sorted(speeches.items()), \
            "Your function makes changes to speeches."
        assert isinstance(result, list), "Result should be a list"
        for item in result:
            assert isinstance(item, tuple) and len(item) == 2, \
                "Each item should be a tuple (name, count)"
            name, count = item
            assert isinstance(name, str), "Name should be a string"
            assert isinstance(count, int), "Count should be an integer"
        assert len(result) == len(solution), "The number of items in the result is not correct"
        assert sorted(name for (name, _) in result) == sorted(name for (name, _) in solution), \
            "The parrot names in the result are not correct"
        assert [name for (name, _) in result] == [name for (name, _) in solution], \
            "The results are not sorted correctly"
        assert sorted(count for (_, count) in result) == sorted(count for (_, count) in solution), \
            "The counts for parrots are not correct"
        assert result == solution, "The result is not correct"
    except AssertionError as e:
        print(f"Test Failed (who_knows_word): {e}")
        print("Test case:")
        pprint((speeches, word))
        print("Your solution:")
        pprint(result)
        print("Correct solution:")
        pprint(solution)
        raise
    print("test OK (who_knows_word)")


def check_pairs_with_shared_vocabulary(speeches, min_shared, solution):
    """
    Check the pairs_with_shared_vocabulary(speeches, min_shared) function.
    """
    speeches_orig = dict(speeches.items())
    result = pairs_with_shared_vocabulary(speeches, min_shared)

    try:
        assert sorted(speeches_orig.items()) == sorted(speeches.items()), \
            "Your function makes changes to speeches."

        assert isinstance(result, dict), "Result should be a dict"

        for key, shared in result.items():
            assert isinstance(key, tuple) and len(key) == 2, \
                "Each key should be a tuple (name1, name2)"
            name1, name2 = key
            assert isinstance(name1, str) and isinstance(name2, str), \
                "Both names in key should be strings"
            assert name1 < name2, "Key should satisfy name1 < name2 alphabetically"

            assert isinstance(shared, set), "Value should be a set of shared words"
            for w in shared:
                assert isinstance(w, str), "Each shared word should be a string"
        assert result.keys() == solution.keys(), "The set of pairs (keys) is not correct"

        for key in result:
            assert result[key] == solution[key], \
                f"Shared vocabulary for pair {key} is not correct"

    except AssertionError as e:
        print(f"Test Failed (pairs_with_shared_vocabulary): {e}")
        print("Test case:")
        pprint((speeches, min_shared))
        print("Your solution:")
        pprint(result)
        print("Correct solution:")
        pprint(solution)
        raise
    print("test OK (pairs_with_shared_vocabulary)")

```

<!-- grade_id: cell-9c0a293e86b2b35e, points: 5, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
def check_function_exists_and_not_notimplemented():
    # 1. speech_to_words(speech)
    assert callable(speech_to_words), "speech_to_words is not defined."
    assert speech_to_words("test") is not NotImplemented, \
        "speech_to_words returned NotImplemented."

    # 2. parrot_vocabularies(speeches)
    assert callable(parrot_vocabularies), "parrot_vocabularies is not defined."
    assert parrot_vocabularies({}) is not NotImplemented, \
        "parrot_vocabularies returned NotImplemented."

    # 3. who_knows_word(speeches, word)
    assert callable(who_knows_word), "who_knows_word is not defined."
    assert who_knows_word({}, "hello") is not NotImplemented, \
        "who_knows_word returned NotImplemented."

    # 4. pairs_with_shared_vocabulary(speeches, min_shared)
    assert callable(pairs_with_shared_vocabulary), \
        "pairs_with_shared_vocabulary is not defined."
    assert pairs_with_shared_vocabulary({}, 1) is not NotImplemented, \
        "pairs_with_shared_vocabulary returned NotImplemented."

    print("test OK (all functions are defined and implemented)")
check_function_exists_and_not_notimplemented()
```

<!-- grade_id: cell-a17b5a7bfc083113, points: 0, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
# Test cases
parrot_speeches_1 = {
    "Max": "Hello Max, hello! Bad Max, Max, Max. Max? Hello!",
}
parrot_speeches_2 = {
    "Pedro": "Hello Max, hello! Bad Max, Max, Max. Max? Hello!",
}
parrot_speeches_3 = {
    "Max": "Hello Max, max Max, hello!  Bad Maxes, Max's Max. Maximus? Hello!",
}
parrot_speeches_4 = {
    "Max": "Hello Max; hello! Bad Max, Max, Max. Max? Hello!",
}
parrot_speeches_5 = {
    "Kim": "Kim, Kim's, Kim. Kims?, kim's! '\"Kim\", kim.",
}
parrot_speeches_6 = {
    "Joe": "Joe! !!! Good Joe! Joejoejoe!",
}

parrot_speeches_7 = {
    "Charlie": "Charlie loves tasty crackers. Crackers, crackers, and more crackers!",
    "Milo": "Milo loves shiny things. Shiny things make Milo happy! Happy Milo, Millo.",
    "Frodo": "Charlie, Max, Milo and Frodo love Crackers!! Frodo loves craskers!",
    "Max": "Max loves tasty crackers. Max? Max! Max loves crackers too. Max, Max",
}

parrot_speeches_8 = {
    "Charlie": "Hello world!",
    "Max": "Hello world!",
    "Anna": "Hello world!"
}

parrot_speeches_9 = {
    "Charlie": "Charlie loves crackers. Crackers are tasty, and Charlie wants more crackers. Hello, Charlie loves tasty crackers!",
    "Max": "Hello Max, hello! Bad Max, Max, Max. Max? Hello!",
    "Carol": "Cracker. Carol's Cracker. Cracker.",
    "Milo": "Milo loves shiny things. Shiny, shiny things make Milo happy. Milo loves sunny days and pretty things.",
    "Kiki": "Kikiki! Kiki-kiki, kikikikikikiki!!!",
    "Polly": "Hello! Hello! Polly, sweet Polly. Tasty crackers, tasty crackers. Hello, Polly wants a sweet cracker! Cracker!",
    "Mia": "Mia, Mia, Mia. Mia, Mia",
}
```

<!-- grade_id: cell-c00debda80c3f76f, points: 10, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
# speech_to_words(speech)
check_speech_to_words(parrot_speeches_9["Max"], ['hello', 'max', 'hello', 'bad', 'max', 'max', 'max', 'max', 'hello'] )
check_speech_to_words(parrot_speeches_9["Charlie"], ['charlie', 'loves', 'crackers', 'crackers', 'are', 'tasty', 'and', 'charlie', 'wants', 'more', 'crackers', 'hello', 'charlie', 'loves', 'tasty', 'crackers'])
check_speech_to_words(parrot_speeches_9["Kiki"], ['kikiki', 'kikikiki', 'kikikikikikiki'])
check_speech_to_words(parrot_speeches_9["Carol"], ['cracker', 'carols', 'cracker', 'cracker'])
check_speech_to_words(parrot_speeches_9["Mia"], ['mia', 'mia', 'mia', 'mia', 'mia'])
check_speech_to_words(parrot_speeches_9["Polly"], ['hello', 'hello', 'polly', 'sweet', 'polly', 'tasty', 'crackers', 'tasty', 'crackers', 'hello', 'polly', 'wants', 'a', 'sweet', 'cracker', 'cracker'])
check_speech_to_words(parrot_speeches_9["Milo"], ['milo', 'loves', 'shiny', 'things', 'shiny', 'shiny', 'things', 'make', 'milo', 'happy', 'milo', 'loves', 'sunny', 'days', 'and', 'pretty', 'things'])

```

<!-- grade_id: cell-08cf66b4a10e5ca6, points: 10, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
# parrot_vocabularies(speech):
check_parrot_vocabularies(parrot_speeches_1, [('Max', {'bad', 'hello', 'max'})])
check_parrot_vocabularies(parrot_speeches_3, [('Max', {'bad', 'hello', 'max', 'maxes', 'maximus', 'maxs'})])
check_parrot_vocabularies(parrot_speeches_5, [('Kim', {'kim', 'kims'})])
check_parrot_vocabularies(parrot_speeches_6, [('Joe', {'good', 'joe', 'joejoejoe'})])
check_parrot_vocabularies(parrot_speeches_7, [('Frodo', {'max', 'craskers', 'charlie', 'crackers', 'frodo', 'and', 'loves', 'love', 'milo'}), ('Milo', {'things', 'make', 'millo', 'shiny', 'loves', 'milo', 'happy'}), ('Charlie', {'charlie', 'crackers', 'more', 'and', 'loves', 'tasty'}), ('Max', {'max', 'too', 'crackers', 'loves', 'tasty'})])
check_parrot_vocabularies(parrot_speeches_8, [('Anna', {'hello', 'world'}), ('Charlie', {'hello', 'world'}), ('Max', {'hello', 'world'})])
check_parrot_vocabularies(parrot_speeches_9, [('Milo', {'sunny', 'pretty', 'things', 'make', 'shiny', 'and', 'loves', 'days', 'milo', 'happy'}), ('Charlie', {'charlie', 'crackers', 'hello', 'more', 'are', 'and', 'loves', 'wants', 'tasty'}), ('Polly', {'hello', 'crackers', 'polly', 'sweet', 'wants', 'cracker', 'tasty', 'a'}), ('Kiki', {'kikikiki', 'kikikikikikiki', 'kikiki'}), ('Max', {'max', 'hello', 'bad'}), ('Carol', {'carols', 'cracker'}), ('Mia', {'mia'})])

```

<!-- grade_id: cell-19fe345f626c9971, points: 10, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
# who_knows_word(speech, word)
check_who_knows_word(parrot_speeches_1, "max", [('Max', 5)])
check_who_knows_word(parrot_speeches_1, "hello", [('Max', 3)])
check_who_knows_word(parrot_speeches_1, "bad", [('Max', 1)])
check_who_knows_word(parrot_speeches_1, "cracker", [])
check_who_knows_word(parrot_speeches_7, "max", [('Max', 6), ('Frodo', 1)])
check_who_knows_word(parrot_speeches_7, "cracker", [])
check_who_knows_word(parrot_speeches_7, "crackers", [('Charlie', 4), ('Max', 2), ('Frodo', 1)])
check_who_knows_word(parrot_speeches_7, "loves", [('Max', 2), ('Charlie', 1), ('Frodo', 1), ('Milo', 1)])
check_who_knows_word(parrot_speeches_7, "happy", [('Milo', 2)])
check_who_knows_word(parrot_speeches_8, "hello", [('Anna', 1), ('Charlie', 1), ('Max', 1)])
check_who_knows_word(parrot_speeches_8, "world", [('Anna', 1), ('Charlie', 1), ('Max', 1)])
check_who_knows_word(parrot_speeches_9, "crackers", [('Charlie', 4), ('Polly', 2)])
check_who_knows_word(parrot_speeches_9, "hello", [('Max', 3), ('Polly', 3), ('Charlie', 1)])
```

<!-- grade_id: cell-79436d0e832d377c, points: 10, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
# check_pairs_with_shared_vocabulary(speeches)
check_pairs_with_shared_vocabulary(parrot_speeches_1, 1, {})
check_pairs_with_shared_vocabulary(parrot_speeches_7, 1, {('Charlie', 'Frodo'): {'charlie', 'and', 'crackers', 'loves'}, ('Charlie', 'Max'): {'tasty', 'crackers', 'loves'}, ('Charlie', 'Milo'): {'loves'}, ('Frodo', 'Max'): {'max', 'crackers', 'loves'}, ('Frodo', 'Milo'): {'milo', 'loves'}, ('Max', 'Milo'): {'loves'}})
check_pairs_with_shared_vocabulary(parrot_speeches_7, 3, {('Charlie', 'Frodo'): {'charlie', 'and', 'crackers', 'loves'}, ('Charlie', 'Max'): {'tasty', 'crackers', 'loves'}, ('Frodo', 'Max'): {'max', 'crackers', 'loves'}})
check_pairs_with_shared_vocabulary(parrot_speeches_7, 4, {('Charlie', 'Frodo'): {'charlie', 'and', 'crackers', 'loves'}})
check_pairs_with_shared_vocabulary(parrot_speeches_8, 1, {('Anna', 'Charlie'): {'hello', 'world'}, ('Anna', 'Max'): {'hello', 'world'}, ('Charlie', 'Max'): {'hello', 'world'}})
check_pairs_with_shared_vocabulary(parrot_speeches_8, 2, {('Anna', 'Charlie'): {'hello', 'world'}, ('Anna', 'Max'): {'hello', 'world'}, ('Charlie', 'Max'): {'hello', 'world'}})
check_pairs_with_shared_vocabulary(parrot_speeches_8, 3, {})
check_pairs_with_shared_vocabulary(parrot_speeches_9, 1, {('Carol', 'Polly'): {'cracker'}, ('Charlie', 'Max'): {'hello'}, ('Charlie', 'Milo'): {'and', 'loves'}, ('Charlie', 'Polly'): {'tasty', 'hello', 'wants', 'crackers'}, ('Max', 'Polly'): {'hello'}})
check_pairs_with_shared_vocabulary(parrot_speeches_9, 2, {('Charlie', 'Milo'): {'and', 'loves'}, ('Charlie', 'Polly'): {'tasty', 'hello', 'wants', 'crackers'}})
check_pairs_with_shared_vocabulary(parrot_speeches_9, 3, {('Charlie', 'Polly'): {'tasty', 'hello', 'wants', 'crackers'}})
check_pairs_with_shared_vocabulary(parrot_speeches_9, 4, {('Charlie', 'Polly'): {'tasty', 'hello', 'wants', 'crackers'}})
check_pairs_with_shared_vocabulary(parrot_speeches_9, 5, {})
```

The following cells contain hidden tests that are evaluated during automatic grading after your submission. The tests check that your solution is sufficiently general.

<!-- grade_id: cell-37d1245db796d390, points: 0, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
```

<!-- grade_id: cell-a3cfe266ff430e72, points: 10, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
```

<!-- grade_id: cell-b0caf3af9c616409, points: 10, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
```

<!-- grade_id: cell-345322b6ad62bec8, points: 10, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
```

<!-- grade_id: cell-d1f1d62ba7fb9830, points: 10, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
```

<!-- grade_id: cell-6c5966bdb2b83529, points: 5, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
```

<!-- grade_id: cell-9fbfcfc1786fe2ef, points: 5, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
```

<!-- grade_id: cell-e7aca2f267644640, points: 5, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
```