# Assignment 11

## Deadline

Wednesday, December 17, 23:59.

This is the **last assignment** from ZPRO 2025 😊

## Introduction

We have a text file `klementinum2023.txt` (generate it using the `generate_files` function below), which contains temperatures measured in 2023 at Klementinum in Prague.
The file has a simple CSV format, where values are separated by commas. The columns have the following meaning:

1. year,
2. month,
3. day,
4. average daily temperature,
5. maximum daily temperature,
6. minimum daily temperature.

An example of one line in the database: `2023,1,1,13.0,15.8,11.0`.

The first line is a comment with column descriptions.

## Task

Implement the function `process_temperatures(in_filename, out_filename, sort='avg')` that performs the following operations on a file named `in_filename` in the above-described format:

* reads the file,
* sorts the data in ascending order by average temperature (`avg`, default), maximum temperature (`max`), or minimum temperature (`min`) (depending on the `sort` parameter),
* saves the data to a file named `out_filename` in the same format (including the same header).

The function returns a status code corresponding to whether the function ran successfully or something failed:

* `0` if there was no error,
* `-1` if there is an issue with the input or output file (does not exist, cannot be created, cannot be read...),
* `-2` if a file format error occurs, e.g., too few or too many values in a line (must be 6), incorrect numeric format (the first three columns are `int` greater than 0 and the next three are `float`).  

If any errors occur, return the error code and delete any newly created file.

## Example

This code

```python
process_temperatures("klementinum2023.txt", "klementinum2023-sorted.txt", "min")
```  

should create the file `klementinum2023-sorted.txt`, where the lines are sorted by minimum temperature in ascending order.
The function returns the value `0`.

## Notes

The function should not have any *side effects*; it should only do what is specified in the task.
Most importantly, do not use the `input` function, do not create other/helper files, do not randomly delete files on the disk, etc.

```python
# Do not delete or edit this cell!
from pathlib import Path


# The following function generates the input files that are used in the public tests.
def generate_files():
    Path("klementinum2023.txt").write_text("""#year,month,day,t_average,t_max,t_min
2023,1,1,13.0,15.8,11.0
2023,1,2,9.4,12.4,8.3
2023,1,3,7.4,11.0,5.4
2023,1,4,7.9,10.7,2.5
2023,1,5,10.5,12.7,9.5
2023,1,6,10.0,12.3,8.1
2023,1,7,7.9,11.5,7.1
2023,1,8,5.7,7.1,4.6
2023,1,9,6.1,9.0,4.8
2023,1,10,5.6,8.0,4.7
2023,1,11,7.1,8.9,2.2
2023,1,12,8.4,9.5,6.2
2023,1,13,9.0,12.1,7.5
2023,1,14,8.1,9.5,6.1
2023,1,15,7.0,11.5,5.4
2023,1,16,5.0,8.2,3.4
2023,1,17,4.1,6.6,3.2
2023,1,18,2.8,4.9,1.7
2023,1,19,0.6,3.4,-0.1
2023,1,20,0.6,2.3,-0.7
2023,1,21,-0.1,1.2,-1.3
2023,1,22,0.9,1.9,0.5
2023,1,23,1.8,2.9,0.8
2023,1,24,2.3,3.3,2.0
2023,1,25,2.0,3.1,2.0
2023,1,26,0.9,2.2,0.9
2023,1,27,0.5,1.3,-0.1
2023,1,28,-0.7,1.5,-1.5
2023,1,29,-1.3,1.1,-2.6
2023,1,30,3.9,5.7,-1.3
2023,1,31,4.4,5.8,2.8
2023,2,1,3.9,7.3,1.6
2023,2,2,5.2,7.3,3.1
2023,2,3,6.4,8.1,1.9
2023,2,4,2.0,7.3,1.6
2023,2,5,-1.0,1.7,-2.3
2023,2,6,-1.8,1.3,-4.3
2023,2,7,-2.5,1.1,-5.9
2023,2,8,-1.5,2.2,-5.3
2023,2,9,-0.3,4.6,-3.6
2023,2,10,0.6,4.4,-2.7
2023,2,11,5.2,6.8,0.5
2023,2,12,7.3,8.6,5.6
2023,2,13,6.3,10.2,5.2
2023,2,14,4.7,6.5,2.8
2023,2,15,1.3,5.0,1.3
2023,2,16,4.0,9.2,-0.6
2023,2,17,11.0,12.4,4.3
2023,2,18,11.6,14.0,10.6
2023,2,19,6.6,11.3,5.1
2023,2,20,10.3,12.7,3.7
2023,2,21,10.4,12.4,7.9
2023,2,22,10.0,12.5,8.2
2023,2,23,9.1,12.9,6.1
2023,2,24,9.3,12.1,8.8
2023,2,25,2.7,8.8,1.4
2023,2,26,0.7,2.9,0.1
2023,2,27,0.4,3.2,-1.8
2023,2,28,1.5,4.5,-1.1
2023,3,1,2.9,7.2,-2.0
2023,3,2,4.4,9.4,-0.7
2023,3,3,1.3,6.1,-0.5
2023,3,4,3.0,6.0,-1.5
2023,3,5,3.4,5.4,2.2
2023,3,6,3.6,6.8,1.2
2023,3,7,5.6,8.2,3.1
2023,3,8,5.6,7.5,4.1
2023,3,9,10.3,13.1,6.4
2023,3,10,10.3,14.5,8.2
2023,3,11,2.0,9.3,2.0
2023,3,12,5.1,7.8,-0.9
2023,3,13,11.9,15.3,7.1
2023,3,14,9.2,14.1,6.8
2023,3,15,3.0,7.9,1.7
2023,3,16,4.1,8.2,-0.6
2023,3,17,8.2,14.0,1.9
2023,3,18,8.6,13.5,3.2
2023,3,19,9.6,13.8,3.9
2023,3,20,11.3,15.7,7.2
2023,3,21,12.7,15.8,8.7
2023,3,22,14.4,19.8,10.1
2023,3,23,15.0,18.9,10.4
2023,3,24,14.2,18.7,12.6
2023,3,25,10.2,13.8,9.1
2023,3,26,9.2,13.5,8.1
2023,3,27,3.1,8.1,2.3
2023,3,28,4.1,7.6,0.5
2023,3,29,7.8,11.6,1.5
2023,3,30,10.0,15.1,7.7
2023,3,31,10.4,13.4,8.9
2023,4,1,11.1,14.7,9.7
2023,4,2,5.1,10.6,4.9
2023,4,3,2.7,5.8,0.6
2023,4,4,3.5,6.3,1.2
2023,4,5,4.1,7.7,-1.1
2023,4,6,5.6,9.8,-0.2
2023,4,7,6.3,9.8,1.6
2023,4,8,6.6,9.5,4.5
2023,4,9,8.2,11.5,4.7
2023,4,10,11.4,16.2,7.4
2023,4,11,9.6,14.1,7.6
2023,4,12,10.9,15.1,2.9
2023,4,13,11.3,13.1,10.0
2023,4,14,6.8,11.1,6.8
2023,4,15,7.8,10.9,6.5
2023,4,16,11.0,15.4,6.1
2023,4,17,9.9,14.8,6.5
2023,4,18,9.9,12.4,8.5
2023,4,19,7.9,11.8,6.0
2023,4,20,10.3,15.8,3.0
2023,4,21,14.6,19.6,8.3
2023,4,22,15.8,21.7,7.4
2023,4,23,15.0,18.3,12.4
2023,4,24,12.9,16.7,10.8
2023,4,25,9.7,13.5,8.2
2023,4,26,8.2,11.1,5.7
2023,4,27,10.5,14.3,5.5
2023,4,28,12.2,17.0,6.9
2023,4,29,13.4,17.3,11.3
2023,4,30,13.0,16.4,10.8
2023,5,1,16.0,20.8,7.0
2023,5,2,14.4,18.2,12.5
2023,5,3,12.4,16.2,8.8
2023,5,4,14.5,19.0,6.3
2023,5,5,18.2,23.3,9.0
2023,5,6,15.8,21.0,14.0
2023,5,7,11.4,15.9,9.7
2023,5,8,12.5,18.6,6.5
2023,5,9,14.7,19.8,7.8
2023,5,10,15.9,21.4,10.0
2023,5,11,18.0,21.5,14.1
2023,5,12,15.8,20.6,11.9
2023,5,13,15.3,17.0,12.3
2023,5,14,14.7,19.4,12.8
2023,5,15,15.4,20.4,11.2
2023,5,16,12.7,17.7,11.0
2023,5,17,11.6,14.5,8.7
2023,5,18,13.0,16.6,8.4
2023,5,19,14.8,18.4,10.9
2023,5,20,18.9,24.9,11.3
2023,5,21,21.7,27.9,14.9
2023,5,22,22.4,27.0,15.9
2023,5,23,17.3,22.0,15.3
2023,5,24,13.3,16.8,11.3
2023,5,25,17.0,23.3,9.2
2023,5,26,17.2,22.1,12.7
2023,5,27,16.8,22.4,11.0
2023,5,28,18.6,23.8,10.8
2023,5,29,19.5,25.3,11.8
2023,5,30,18.3,23.3,14.4
2023,5,31,19.8,26.4,11.3
2023,6,1,20.4,28.7,11.7
2023,6,2,16.7,21.1,13.2
2023,6,3,16.9,21.8,11.4
2023,6,4,18.8,25.6,11.0
2023,6,5,15.4,22.9,14.3
2023,6,6,17.3,20.6,15.0
2023,6,7,20.7,24.6,16.4
2023,6,8,21.2,26.3,16.7
2023,6,9,19.5,24.7,17.1
2023,6,10,21.8,26.4,15.9
2023,6,11,21.6,26.3,14.8
2023,6,12,20.4,24.1,15.7
2023,6,13,17.5,23.4,12.4
2023,6,14,17.7,21.6,11.6
2023,6,15,17.5,23.6,11.7
2023,6,16,17.3,23.0,12.4
2023,6,17,17.4,22.7,14.7
2023,6,18,22.0,27.7,12.5
2023,6,19,24.3,29.8,16.3
2023,6,20,26.4,33.4,17.9
2023,6,21,25.7,31.6,22.6
2023,6,22,27.6,32.4,19.1
2023,6,23,20.5,28.2,18.9
2023,6,24,21.1,25.9,15.8
2023,6,25,23.2,28.4,15.2
2023,6,26,24.7,33.1,17.5
2023,6,27,18.4,23.6,16.7
2023,6,28,19.9,24.3,14.7
2023,6,29,23.1,28.7,13.2
2023,6,30,20.5,24.9,18.9
2023,7,1,23.2,27.3,17.9
2023,7,2,22.2,25.7,19.2
2023,7,3,22.4,27.4,16.0
2023,7,4,22.9,28.1,16.2
2023,7,5,22.1,28.2,16.8
2023,7,6,21.9,28.1,14.8
2023,7,7,24.2,31.3,14.8
2023,7,8,26.3,32.6,18.0
2023,7,9,27.7,34.3,18.1
2023,7,10,24.6,31.9,20.3
2023,7,11,27.3,33.1,19.8
2023,7,12,26.8,31.8,21.0
2023,7,13,23.7,28.1,19.7
2023,7,14,25.1,30.5,17.6
2023,7,15,30.1,37.6,18.8
2023,7,16,24.3,31.5,21.8
2023,7,17,26.0,31.3,20.9
2023,7,18,25.2,31.3,18.4
2023,7,19,24.3,30.6,18.7
2023,7,20,21.2,25.5,19.0
2023,7,21,19.6,26.4,16.1
2023,7,22,21.2,26.0,15.9
2023,7,23,25.0,30.3,17.4
2023,7,24,24.2,28.9,19.6
2023,7,25,18.9,24.5,16.7
2023,7,26,17.6,20.8,14.0
2023,7,27,18.2,23.9,13.0
2023,7,28,22.6,26.7,17.2
2023,7,29,20.3,26.0,19.6
2023,7,30,20.1,25.1,17.6
2023,7,31,20.9,24.6,16.2
2023,8,1,20.4,23.4,18.1
2023,8,2,19.7,23.9,16.7
2023,8,3,21.6,26.0,17.8
2023,8,4,21.7,24.9,18.6
2023,8,5,17.3,21.6,16.4
2023,8,6,16.9,20.7,15.7
2023,8,7,16.1,20.1,14.2
2023,8,8,18.9,22.7,13.0
2023,8,9,17.0,21.3,14.2
2023,8,10,18.6,23.3,10.8
2023,8,11,22.6,27.4,15.8
2023,8,12,24.8,31.5,16.4
2023,8,13,25.1,29.0,20.9
2023,8,14,27.6,32.3,20.9
2023,8,15,27.2,33.5,21.5
2023,8,16,23.9,31.3,20.2
2023,8,17,24.1,29.5,18.9
2023,8,18,24.9,29.8,19.1
2023,8,19,26.8,32.1,19.9
2023,8,20,27.4,32.1,21.5
2023,8,21,27.2,31.6,22.1
2023,8,22,26.6,33.1,21.7
2023,8,23,25.4,29.5,21.6
2023,8,24,25.7,31.3,18.8
2023,8,25,25.6,30.7,21.7
2023,8,26,23.3,28.6,20.9
2023,8,27,18.9,21.8,18.8
2023,8,28,14.2,18.9,13.0
2023,8,29,14.7,16.7,12.1
2023,8,30,16.9,19.7,15.0
2023,8,31,17.3,20.9,14.0
2023,9,1,18.0,21.0,13.9
2023,9,2,20.4,23.3,16.8
2023,9,3,19.1,23.3,17.0
2023,9,4,19.2,23.7,13.6
2023,9,5,19.7,25.4,14.0
2023,9,6,20.5,26.7,13.1
2023,9,7,22.2,27.9,15.1
2023,9,8,22.4,28.4,15.9
2023,9,9,21.9,28.9,14.7
2023,9,10,22.0,28.5,15.1
2023,9,11,23.2,29.6,16.1
2023,9,12,23.9,29.9,18.2
2023,9,13,22.2,26.8,18.7
2023,9,14,17.9,21.8,17.1
2023,9,15,18.4,23.2,12.4
2023,9,16,20.9,26.6,14.7
2023,9,17,21.3,27.2,14.2
2023,9,18,21.0,28.6,16.7
2023,9,19,19.6,24.0,17.8
2023,9,20,20.1,26.0,14.2
2023,9,21,19.7,24.3,15.2
2023,9,22,16.4,20.1,14.9
2023,9,23,15.5,18.7,13.5
2023,9,24,14.2,18.8,9.7
2023,9,25,16.2,21.2,10.0
2023,9,26,18.6,24.4,12.8
2023,9,27,19.6,26.0,14.5
2023,9,28,18.5,25.1,12.9
2023,9,29,20.2,25.8,12.7
2023,9,30,16.3,21.8,15.6
2023,10,1,17.2,22.7,10.1
2023,10,2,18.6,25.8,12.2
2023,10,3,18.8,27.8,12.5
2023,10,4,15.4,19.5,12.5
2023,10,5,14.0,18.5,9.7
2023,10,6,14.7,19.4,10.1
2023,10,7,17.6,20.5,13.9
2023,10,8,12.3,18.7,11.3
2023,10,9,14.8,16.6,11.3
2023,10,10,17.6,21.4,14.3
2023,10,11,19.6,25.2,14.2
2023,10,12,17.8,20.7,15.0
2023,10,13,19.8,25.5,15.8
2023,10,14,14.8,22.5,13.5
2023,10,15,9.0,14.0,7.4
2023,10,16,7.8,11.0,4.1
2023,10,17,7.3,11.3,4.2
2023,10,18,8.1,12.8,2.7
2023,10,19,10.3,12.6,6.7
2023,10,20,14.1,16.6,10.3
2023,10,21,16.5,20.9,11.6
2023,10,22,14.1,17.4,12.5
2023,10,23,12.8,17.2,7.6
2023,10,24,12.3,14.5,11.0
2023,10,25,12.9,16.0,11.1
2023,10,26,13.0,15.6,11.9
2023,10,27,11.4,13.4,11.1
2023,10,28,12.5,15.4,10.5
2023,10,29,13.9,18.2,10.7
2023,10,30,14.0,17.7,10.1
2023,10,31,12.6,14.5,11.9
2023,11,1,12.8,16.7,11.2
2023,11,2,11.2,14.4,8.3
2023,11,3,9.2,11.7,8.3
2023,11,4,9.8,12.5,7.0
2023,11,5,12.4,15.4,9.6
2023,11,6,11.3,14.3,9.9
2023,11,7,10.4,13.5,8.3
2023,11,8,8.3,11.2,7.4
2023,11,9,9.0,13.0,4.7
2023,11,10,8.2,11.8,7.4
2023,11,11,5.8,8.0,4.1
2023,11,12,6.2,10.3,5.3
2023,11,13,7.1,8.8,4.6
2023,11,14,11.6,13.4,8.8
2023,11,15,10.2,12.3,9.6
2023,11,16,8.1,11.0,7.3
2023,11,17,5.8,7.7,4.8
2023,11,18,5.1,7.2,4.4
2023,11,19,9.1,12.0,3.4
2023,11,20,10.7,13.0,9.2
2023,11,21,7.0,10.9,4.8
2023,11,22,1.6,4.8,1.3
2023,11,23,6.3,8.1,2.0
2023,11,24,4.9,9.2,3.2
2023,11,25,2.6,4.5,1.5
2023,11,26,1.5,3.2,0.1
2023,11,27,3.4,5.5,1.9
2023,11,28,0.3,3.2,-0.6
2023,11,29,0.5,1.6,-0.9
2023,11,30,1.5,3.9,0.2
2023,12,1,-0.4,1.3,-0.5
2023,12,2,-2.0,0.1,-2.5
2023,12,3,-1.5,-0.2,-2.9
2023,12,4,-3.3,-0.9,-6.1
2023,12,5,-0.9,1.3,-2.2
2023,12,6,0.9,2.8,-1.9
2023,12,7,2.6,3.6,1.3
2023,12,8,0.3,2.8,-1.6
2023,12,9,3.2,4.3,0.8
2023,12,10,5.9,8.1,2.2
2023,12,11,8.6,10.4,6.1
2023,12,12,8.1,10.0,7.3
2023,12,13,7.4,10.3,6.5
2023,12,14,6.1,8.5,5.3
2023,12,15,5.3,6.5,3.9
2023,12,16,6.7,8.4,5.4
2023,12,17,6.4,9.7,4.3
2023,12,18,7.6,11.1,5.0
2023,12,19,8.3,11.3,5.3
2023,12,20,6.2,8.4,5.4
2023,12,21,6.4,9.9,5.6
2023,12,22,2.4,6.5,1.4
2023,12,23,3.3,4.0,1.3
2023,12,24,9.8,10.8,4.0
2023,12,25,11.2,12.2,10.6
2023,12,26,9.9,11.4,8.9
2023,12,27,6.5,9.2,5.3
2023,12,28,8.5,11.1,2.4
2023,12,29,10.7,12.1,8.7
2023,12,30,8.4,11.0,7.4
2023,12,31,4.7,7.4,2.6""")
    Path("kl_err1a.txt").write_text("""#year,month,day,t_average,t_max,t_min
2023,10,1,1.5,5.0,0.2
2023,10,11,aaa,1.5,1.4""")
    Path("kl_err1b.txt").write_text("""#year,month,day,t_average,t_max,t_min
2023,10,1,5,6.5,1.2
2023,10,x,1.2,1.3,1.4""")
    Path("kl_err2a.txt").write_text("""#year,month,day,t_average,t_max,t_min
2023,12,29,10.7,12.1,8.7
2023,12,28,10.7,12.1""")
    Path("kl_err2b.txt").write_text("""#year,month,day,t_average,t_max,t_min
2023,12,29,10.7,12.1,8.7
2023,12,28,10.7,12.1,8.7,100.0""")
    Path("kl_err3a.txt").write_text("""#year,month,day,t_average,t_max,t_min
2023,12.1,29,10.7,12.1,8.7
2023,12,28,10.7,12.1,8.7""")
    Path("kl_err3b.txt").write_text("""#year,month,day,t_average,t_max,t_min
2023,12,29,10.7,12.1,8.7
2023,12,28,10.7,0xff,8.7""")
    Path("kl_err4a.txt").write_text("""#year,month,day,t_average,t_max,t_min
2023,-12,29,10.7,12.1,8.7
2023,12,28,10.7,12.1,8.7""")
    Path("kl_err4b.txt").write_text("""#year,month,day,t_average,t_max,t_min
2023,12,29,10.7,12.1,8.7
2023,12,-28,10.7,15.0,8.7""")


generate_files()
```

<!-- grade_id: cell-a9846e190b320832, points: 0, grade: false, solution: true -->
```python
#
import os

#! Process Temperatures Function
def process_temperatures(in_filename, out_filename, sort="avg"):
     fileExisted = os.path.exists(out_filename)

     #* Rewrite document into array
     try:
          with open(f"{in_filename}", "rt") as input:
               lines = input.readlines()
     except:
          return -1

     #* Clean & check for inconsistencies
     header = lines.pop(0)
     fixedLines = []

     for line in lines:
          data = line.split(",")

          if len(data) != 6:
               return -2
          
          try:
               data[0] = int(data[0])
               data[1] = int(data[1])
               data[2] = int(data[2])
               data[3] = float(data[3])
               data[4] = float(data[4])
               data[5] = float(data[5].replace("\n", ""))
          except:
               return -2

          if data[0] < 0 or data[1] < 0 or data[1] > 12 or data[2] < 0 or data[2] > 31:
               return -2

          fixedLines.append(data)

     #* Sort accordingly
     sortingIndex = -1

     if sort == "avg":
          sortingIndex = 3
     elif sort == "max":
          sortingIndex = 4
     elif sort == "min":
          sortingIndex = 5
     else:
          return -2
     
     fixedLines.sort(key=lambda x: x[sortingIndex])

     #* Build a string
     outputArray = []
     for line in fixedLines:
          lineString = str(line[0]) + "," + str(line[1]) + "," + str(line[2]) + "," + str(line[3]) + "," + str(line[4]) + "," + str(line[5])
          outputArray.append(lineString)
     outputString = header + "\n".join(outputArray)

     #* Write to output file
     try:
          with open(out_filename, "wt") as output:
               output.write(outputString)
               return 0
     except:
          if not fileExisted and os.path.exists(out_filename):
               os.remove(out_filename)
          return -1
```

# Tests
The following cells contain public tests that you can use for basic validation of your solution. **Click on the validate button before submitting!**

<!-- grade_id: cell-33a51e495c27a0f3, points: 15, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
# Tests for basic function:
from pathlib import Path

# klementinum2023-sorted-1.txt by mel obsahovat stejna data, ale serazena podle prumerne teploty
assert process_temperatures("klementinum2023.txt", "klementinum2023-sorted-1.txt") == 0
assert Path("klementinum2023-sorted-1.txt").exists()

# klementinum2023-sorted-1a.txt by mel obsahovat stejna data, ale serazena podle prumerne teploty
assert process_temperatures("klementinum2023.txt", "klementinum2023-sorted-1a.txt", "avg") == 0
assert Path("klementinum2023-sorted-1a.txt").exists()

# klementinum2023-sorted-2.txt by mel obsahovat stejna data, ale serazena podle maximalni teploty vzestupne
assert process_temperatures("klementinum2023.txt", "klementinum2023-sorted-2.txt", "max") == 0
assert Path("klementinum2023-sorted-2.txt").exists()

# klementinum2023-sorted-3.txt by mel obsahovat stejna data, ale serazena podle minimalni teploty vzestupne
assert process_temperatures("klementinum2023.txt", "klementinum2023-sorted-3.txt", "min") == 0
assert Path("klementinum2023-sorted-3.txt").exists()

# klementinum2023-sorted-4.txt by mel obsahovat stejna data, ale serazena podle prumerne teploty
# ze souboru serazeneho ze souboru serazeneho podle minima
assert process_temperatures("klementinum2023-sorted-3.txt", "klementinum2023-sorted-4.txt") == 0
assert Path("klementinum2023-sorted-4.txt").exists()

print("All tests passed.")
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[35], line 6
      3 from pathlib import Path
      5 # klementinum2023-sorted-1.txt by mel obsahovat stejna data, ale serazena podle prumerne teploty
----> 6 assert process_temperatures("klementinum2023.txt", "klementinum2023-sorted-1.txt") == 0
      7 assert Path("klementinum2023-sorted-1.txt").exists()
      9 # klementinum2023-sorted-1a.txt by mel obsahovat stejna data, ale serazena podle prumerne teploty

AssertionError:
```

<!-- grade_id: cell-0cbf8e801ce7a7f7, points: 15, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
# Test for missing file:
from pathlib import Path

# soubor abc.txt neexistuje
assert process_temperatures("abc.txt", "error.txt") == -1
assert Path("error.txt").exists() is False

# read only soubor
Path("readonly.txt").touch()
Path("readonly.txt").chmod(0o444)
assert process_temperatures("klementinum2023.txt", "readonly.txt") == -1

# wrong path
assert process_temperatures("klementinum2023.txt", "A:/floppy.txt") == -1

print("All tests passed.")
```

<!-- grade_id: cell-582ddcb5090415ab, points: 10, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
# Test for invalid format (temperature, date)

# misto teploty je text
assert process_temperatures('kl_err1a.txt', 'error.txt') == -2
assert Path('error.txt').exists() is False

# misto data je text
assert process_temperatures('kl_err1b.txt', 'error.txt') == -2
assert Path('error.txt').exists() is False

print("All tests passed.")
```

<!-- grade_id: cell-86d5a1660bccd97a, points: 10, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
# Test for invalid format (date > 0, date = int)

# misto data je float
assert process_temperatures('kl_err3a.txt', 'error.txt') == -2
assert Path('error.txt').exists() is False

# misto data je hex
assert process_temperatures('kl_err3b.txt', 'error.txt') == -2
assert Path('error.txt').exists() is False

# datum < 0
assert process_temperatures('kl_err4a.txt', 'error.txt') == -2
assert Path('error.txt').exists() is False

# datum < 0
assert process_temperatures('kl_err4b.txt', 'error.txt') == -2
assert Path('error.txt').exists() is False

print("All tests passed.")
```

<!-- grade_id: cell-86d5a1660bccd97b, points: 10, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
# Test for invalid format (not enough data)

# chybi jedna hodnota
assert process_temperatures("kl_err2a.txt", "error.txt") == -2
assert Path("error.txt").exists() is False

# prebyva jedna hodnota
assert process_temperatures("kl_err2b.txt", "error.txt") == -2
assert Path("error.txt").exists() is False

print("All tests passed.")
```

The following cells contain hidden tests that are evaluated during automatic grading after your submission. The tests check that your solution is sufficiently general.
**Correct filtering and sorting of items in the output file is checked only by the hidden tests!**

<!-- grade_id: cell-6c5966bdb2b83529, points: 20, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
```

<!-- grade_id: cell-e213c6c1b630d404, points: 20, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
```