# Assignment 1

This is the first assignment from ZPRO, the following assignments will be released weekly.
The total number of assignments will be 11 and you can get 100 points for each.
The first assignment also contains detailed instructions how to manipulate with assignment files that should be followed for the rest of the assignments.

## Deadline

Sunday, October 5, 23:59.

# Guideline

This _Jupyter notebook_, which contains the task assignment and soon will contain your solution, is __special__, unlike regular notebooks that we use during lectures.
Some cells contain code for _partial validation_ of your solution to the task, while other cells then contain _hidden code_ for additional validation after submission.
On the technical level, evaluation of your assignments uses _hidden attributes_ in each cell, which are compared with the original assignment and their damage can break the validation process.

Unfortunately, the software used for submitting solutions is not very robust and you have to follow these rules when editing this notebook:

- __Edit only the cells that are intended for it.__
  The solution of the assignment has to be added in a cell that contains a placeholder `YOUR CODE HERE`.
  Typically, this will be just one cell, but later assignments may be divided into several parts across multiple cells.
- __Do not leave__ the command `raise NotImplementedError()` in the edited cells – it needs to be __replaced__ with your solution.
- __Do not add new cells__ and do not __move__ existing cells.
  If you need more space for thinking and testing your solution, copy the code into another notebook.
  If you _duplicate a cell and delete it_, then there is a risk that you will accidentally remove the original one, which would change its ID in hidden attributes and break automatic validation.
- When submitting an assignment, all files in the folder containing this notebook are submitted.
  __Do not add__ any other files (helper notebooks) and __do not rename__ the notebook with the assignment.
- Save the file before submitting it to make sure that the latest version of the file is submitted.

If you follow these rules and no error is displayed after clicking on the _Validate_ button, then you can submit your solution for full validation.

If you find a better solution before the deadline, you may resubmit it – only the last submission in time will be taken into account.

# Auxilary cell

The following cell contains auxiliary code that is needed for the task validation. **Do not do anything with this cell!**

```python
# this is a read-only cell - do not edit, delete, or move!
import zprotests
```

# Task assignment

Write a program that prints the number of __bytes__ needed to store the text "Hobbyte" in a binary representation using the ASCII table.
The program should not do anything else - especially, it must not print anything else.

<!-- grade_id: cell-d46e4b44d3d19218, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()
string = "Hobbyte"
length = len(string)
print(length)
```

Output:
```
7
```

## Tests

The following cells contain public or hidden tests for automatic validation of the assignment.
**Click on the Validate button before submitting** to evaluate the public tests.
Hidden tests are evaluated only after submission and you will see them only in the feedback.
All tests affect your score.

<!-- grade_id: cell-c75df1d93e7f7bbb, points: 20, grade: true, solution: false -->
```python
# this is a read-only cell used for automatic validation - do not edit, delete, or move!
assert zprotests.stdout
```

<!-- grade_id: cell-a3a55b3089fd00d7, points: 20, grade: true, solution: false -->
```python
# this is a read-only cell used for automatic validation - do not edit, delete, or move!
```

<!-- grade_id: cell-afd8a44c5e345b90, points: 20, grade: true, solution: false -->
```python
# this is a read-only cell used for automatic validation - do not edit, delete, or move!
```

<!-- grade_id: cell-fd65dbc774726d25, points: 10, grade: true, solution: false -->
```python
# this is a read-only cell used for automatic validation - do not edit, delete, or move!
```

<!-- grade_id: cell-28fd6851021bc365, points: 10, grade: true, solution: false -->
```python
# this is a read-only cell used for automatic validation - do not edit, delete, or move!
```

<!-- grade_id: cell-5cb2985d4b80c007, points: 20, grade: true, solution: false -->
```python
# this is a read-only cell used for automatic validation - do not edit, delete, or move!
```

```python

```