# Assignment 3

## Deadline

Wednesday, October 22, 23:59.

### Task
We consider the following ([Collatz](https://en.wikipedia.org/wiki/Collatz_conjecture)) sequence:

- Take an integer $ n $. As long as $ n > 1 $, update its value as follows: If $ n $ is divisible by two, divide it by two.
 Otherwise, multiply it by three and add one.
- Once the sequence reaches the value $1$, that term is considered the **last**.
- The given integer $ n >= 1 $ is taken as the **first term** of the sequence.  
- If the user enters a value $n < 1$, the sequence has **no terms**.


Write a program that reads two integers:
1. $ n $ – the first term of the sequence,
2. $ k $ – which **even** term of the sequence we are interested in,

and then computes and prints three values:

1. **X** – the **largest odd number** in the given Collatz sequence. If the sequence has no terms, X will be `-1`.
2. **Y** – the **number of terms** in the Collatz sequence that lie within the interval `[10, 100)` ($10$ inclusive, $100$ exclusive).
3. **Z** – the $k$-th **even** term of the Collatz sequence.     If the sequence does not contain at least $k$ even terms, or if $k < 1$, Z will be `-1`.

### Important notes

- Use the `input` function to read the two integers $ n $  and $ k $.  For automatic evaluation, call `input` **exactly twice** – first for $ n $, then for $ k $.
- Print the results using the `print` function **once only**, on a single line in the format  `X Y Z` (three integers separated by a single space). Do not print anything else.

For example, for the input values $n = 12$ and $k = 4$ the program should output `5 3 16`.
For $n = 13$ and $k = 1$ the program should output `13 5 40`. For $n = 1$ and $k = 1$ the program should output `1 0 -1`. The correctness of the output for these three inputs is verified by **public tests**. 

*Final tip:* don’t forget to *save* your final solution and click the *Validate* button before *submitting*.

```python
# this is a read-only cell - do not edit, delete, or move!
import zprotests
```

<!-- grade_id: cell-28dbe96b036581b3, points: 0, grade: false, solution: true -->
```python
# Collatz Conjecture ++
## Input
n = int(input())
k = int(input())

## Collatz sequence
def getCollatz(number):
    array = []
    
    if number < 1:
        return []

    while number > 1:
        array.append(number)

        if (number % 2) == 0:
            number = number // 2
        else:
            number = (number * 3) + 1
    array.append(1)
    return array

## X: largest odd number
def getLargestOddNumber(inputArray):
    array = inputArray.copy() ### python only & doesn't copy automatically

    if len(array) != 0:
        oddArray = []

        for i in array:
            if (i % 2) != 0:
                oddArray.append(i)
        return max(oddArray)
    else:
        return -1

## Y: # numbers in [10,100)
def getIntervalNumber(inputArray):
    array = inputArray.copy()
    count = 0

    for i in array:
        if (i >= 10) and (i < 100): ### i is a value of array[j]
            count += 1
    return count

## Z: # k-th even term
def getSomeEvenNumber(inputArray, k):
    array = inputArray.copy()
    evenArray = []

    if k < 1:
        return -1

    for i in array:
        if (i % 2) == 0:
            evenArray.append(i)

    if k <= len(evenArray):
        return evenArray[k - 1]
    else:
        return -1

## Output
collatzArray = getCollatz(n)

x = getLargestOddNumber(collatzArray)
y = getIntervalNumber(collatzArray)
z = getSomeEvenNumber(collatzArray, k)

print(f"{x} {y} {z}")
```

Output:
```
-1
 1
-1 0 -1
```

<!-- grade_id: cell-eca876ce26b58809, points: 5, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!

# This cell validates the example values given in the text of the assignment.
# If something fails and you think that the result is correct, check that you
# use the input and print functions as required by the assignment.

with zprotests.mock_input(inputs=[12, 4]):
    %rerun_quiet -g 'input('

assert zprotests.stdout == "5 3 16\n"
```

<!-- grade_id: cell-700f2e2b81a58380, points: 5, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!

# This cell validates the example values given in the text of the assignment.
# If something fails and you think that the result is correct, check that you
# use the input and print functions as required by the assignment.

with zprotests.mock_input(inputs=[13, 1]):
    %rerun_quiet -g 'input('

assert zprotests.stdout == "13 5 40\n"
```

<!-- grade_id: cell-0f004e450faa3b82, points: 5, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!

# This cell validates the example values given in the text of the assignment.
# If something fails and you think that the result is correct, check that you
# use the input and print functions as required by the assignment.

with zprotests.mock_input(inputs=[1, 1]):
    %rerun_quiet -g 'input('

assert zprotests.stdout == "1 0 -1\n"
```

<!-- grade_id: cell-b58b997e4fc8eb59, points: 10, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!

```

<!-- grade_id: cell-4f994da053aefa6f, points: 10, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!

```

<!-- grade_id: cell-bfc2bec0c56fb157, points: 10, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!

```

<!-- grade_id: cell-67e66bb49251ab13, points: 10, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!

```

<!-- grade_id: cell-4f658b8d97b9a56c, points: 10, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!

```

<!-- grade_id: cell-162fee1aa8057ad2, points: 10, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!

```

<!-- grade_id: cell-71dd663b3d4bee2a, points: 10, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!

```

<!-- grade_id: cell-c913161a658852ab, points: 5, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!

```

<!-- grade_id: cell-e9b776e6580c463d, points: 5, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!

```

<!-- grade_id: cell-f6fe25ac5c242630, points: 5, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!

```