# Assignment 4

## Deadline

Wednesday, October 29, 23:59.

## Task

The sum of all prime numbers greater than $1$ and less than $10$ is $2 + 3 + 5 + 7 = 17$.

Write a function named `sum_primes`, which takes two parameters $a$ and $b$, where the latter will be optional and it will be possible to use the function in two ways:

1. If only the first parameter $a$ is given, function `sum_primes` will return a sum of all prime numbers in interval $[1, a)$ ($a$ is not included).
2. If both parameters $a$ and $b$ are given, function `sum_primes` will return a sum of all prime numbers in interval $[a, b)$ ($a$ is included, $b$ is not included).

**Comments and hints:**

- The number $1$ is not a prime number.
- The function should always return an integer number.
- The function should return $0$ if the interval is empty or if $a > b$.
- For clarity of the code define a helper function `is_prime` and use it in the function `sum_primes`.
- All tests use only positive integers, i.e. no need to check for negative numbers or non-integers.
- **Do not use the function `input` in this assignment.**

<!-- grade_id: cell-28dbe96b036581b3, points: 0, grade: false, solution: true -->
```python
# Why is something called "Comments and hints" required to pass the tests? Hint implies non-mandatory

# Sum Interval Of Primes
## Find primes helper f(x)
def is_prime(n):
    if n == 0 or n==1:
        return False

    sqrt = int(n**(1/2))

    for j in range (2, sqrt + 1):
        if n % j == 0:
            break
    else:
       return True
    return False

def findPrimes(down, up):
    primes = []

    for i in range(down, up):
        if is_prime(i) == True:
            primes.append(i)
    return primes

## Sum primes main f(x)
def sum_primes(a, b=0):
    primes = []
    
    if b == 0:
        primes = findPrimes(1, a)
        
    elif a <= b:
        primes = findPrimes(a, b)
    else:
        return 0
    return sum(primes)
```

The following cells contain public tests that you can use for basic validation of your solution. **Click on the validate button before submitting!**

```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
import ipytest
ipytest.autoconfig(raise_on_error=True, addopts=["-p", "no:cacheprovider"])  # type: ignore
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[9], line 2
      1 # This is a read-only cell used for automatic validation - do not edit, delete, or move!
----> 2 import ipytest
      3 ipytest.autoconfig(raise_on_error=True, addopts=["-p", "no:cacheprovider"])  # type: ignore

ModuleNotFoundError: No module named 'ipytest'
```

<!-- grade_id: cell-5443d33a7c80dc12, points: 5, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
assert "sum_primes" in globals(), "The function 'sum_primes' is not defined."
assert sum_primes(0) == 0, "The function 'sum_primes' does not accept one parameter or returns wrong result for zero."
assert sum_primes(0, 0) == 0, "The function 'sum_primes' does not accept two parameters or returns wrong result for two zeros."
```

<!-- grade_id: cell-799821c89e9f18c0, points: 5, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
assert "is_prime" in globals(), "The function 'is_prime' is not defined."
assert is_prime(0) is False, "The function 'is_prime' does not work for number 0."
```

<!-- grade_id: cell-b4eeb6797b63d39d, points: 5, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
assert sum_primes(10) == 17
assert sum_primes(11) == 17
assert sum_primes(12) == 28
assert sum_primes(13) == 28
assert sum_primes(1, 10) == 17
assert sum_primes(2, 10) == 17
assert sum_primes(3, 10) == 15
assert sum_primes(4, 10) == 12
```

<!-- grade_id: cell-74cc1e48d19a1fc9, points: 5, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
assert is_prime(1) is False
assert is_prime(2) is True
assert is_prime(3) is True
assert is_prime(4) is False
assert is_prime(5) is True
assert is_prime(6) is False
assert is_prime(7) is True
```

The following cells contain hidden tests that are evaluated during automatic grading after your submission. The tests check that your functions are sufficiently general.

<!-- grade_id: cell-539179e81fa8a20f, points: 50, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
```

<!-- grade_id: cell-1cec430b2fffaa3e, points: 30, grade: true, solution: false -->
```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
```