# Assignment 8

## Deadline

Wednesday, November 26, 23:59.

# Assignment

Your task is to write a function that determines whether a given input string can represent a [Czech passenger-car license plate (SPZ)](https://en.wikipedia.org/wiki/Vehicle_registration_plates_of_the_Czech_Republic) valid under the rules of 2025.

![L0VEZPR0.png](L0VEZPR0.png)

In this task, we distinguish only four types of license plates:
1. __standard__ – standard passenger car license plate
2. __custom__ – personalized license plate (so-called _"SPZ na přání"_)
3. __electric__ – license plate for electric vehicles
4. __invalid__ – anything that does not fall into any of the above categories

### Preprocessing of the input string:

Before analyzing the license plate:

- Remove spaces or any whitespace characters at the beginning or end of the input string.
- The first three characters in the input string may be followed by a single space (e.g., "L0V EZPR0"), but this is not required (i.e., it can also be "L0VEZPR0"). If such a space is present, remove it completely; it is not considered part of the license plate.
- Spaces anywhere else are not allowed. If present, the plate is invalid.

### Universal rules (apply for all valid license plates):

A valid Czech license plate:

- consists only of digits `0`-`9` and uppercase Latin letters without diacritics;
- must not contain the letters `G`, `O`, `Q`, and `W` (because they are easily confused);
- must contain at least one digit (purely alphabetical strings are not allowed);
- must not contain special characters (e.g., `-`, `*`, `/`, `?`, `!`, `§`, etc.);
- must not contain whitespace (except as described in the preprocessing section above).

The Czech digraph `CH` is not a single letter; it counts as two separate Latin letters.

### Standard passenger car license plate:

A standard licence plate:

- contains exactly 7 characters after preprocessing.
- has the format `DLCDDDD`, where `D` stands for a digit, `L` is any allowed letter, `C` is any allowed character (letter or digit);
- must not start with the digit `0`;
- the second character is a __region code__ (see the [list on the ministry website](https://md.gov.cz/Dokumenty/Silnicni-doprava/Registrace-vozidel/Kody-kraju-registracni-znacky?returl=/Dokumenty/Silnicni-doprava/Registrace-vozidel)).

__Special rule for Prague__ (valid since 2023): If the region code corresponds to Prague, then the __fourth position__ may contain __a letter__ instead of a digit.

### Custom license plate:

A personalized (custom) license plate:

- contains 7 or 8 characters after preprocessing;
- satisfies all universal rules;
- but does not satisfy at least one rule for standard plates;
- must not start with the prefix `EL` (those are reserved for electric vehicles).

### License plate for electric vehicles:

An electric-vehicle license plate:

- contains exactly 7 characters after preprocessing;
- begins with the prefix `EL`;
- has the format: `ELDDDCC` (prefix `EL`, followed by three digits, ending with any two allowed characters).

### Invalid license plate
A license plate is considered invalid if it:
- violates universal rules;
- or does not match the structure of a standard, custom, or electric vehicle plate;
- or belongs to a category we are not distinguishing (e.g., diplomatic, historic, or special vehicles).

## Input and output format
Write a function `analyze_license_plate(plate)` that has a single input parameter of type `str`.

The function always returns a single value of type `str`:
- `"standard"` if the input string is a standard passenger car license plate;
- `"custom"` if the input string is a custom license plate;
- `"electric"` if the input string is a license plate for electric vehicles;
- `"invalid"` if none of the previous cases apply.

### Implementation guidelines 
- Do not use external libraries (only standard Python).
- Do not use global variables.
- Do not use `input()` in the final solution (input data is provided automatically).  
- Do not use `print()` in the final solution.  
- Create small helper functions as needed.
- For sample input/output data, refer to the public test cases section. However, note that the public tests *do not cover all special cases*, so pay attention to the assignment.

<!-- grade_id: cell-a9846e190b320832, points: 0, grade: false, solution: true -->
```python
#

#! Imports
from re import fullmatch, search

#! Helper Functions
#* Preprocessing
def preprocessPlate(plate):
     purePlate = plate.strip()

     # get rid of space on pos 4
     if len(purePlate) >= 4:
          halfPlate = purePlate[0:3]
          if purePlate[3] == " ":
               secondHalfPlate = purePlate[4:]
          else:
               secondHalfPlate = purePlate[3:]          
          return halfPlate + secondHalfPlate
     return purePlate

#* Plate is valid
def isValid(seq):
     onlyDigitsAndUppercase = r'([0-9]|[A-Z])+'
     forbiddenLetters = r'G|O|Q|W'
     digit = r'[0-9]'

     if fullmatch(onlyDigitsAndUppercase, seq) and not search(forbiddenLetters, seq) and search(digit, seq):
          return True
     return False

#* Plate is standard
def isStandard(seq):
     standard = r'[1-9][SULKHEPCJBMTZ][0-9A-Z][0-9]{4}'
     pragueStandard = r'[1-9]A[0-9A-Z]{2}[0-9]{3}'
     if fullmatch(standard, seq) or fullmatch(pragueStandard, seq):
          return True
     return False

def isCustom(seq):
     custom = r'[0-9A-Z]{7,8}'
     leadingEl = r'EL[0-9A-Z]{5,6}'
     if fullmatch(custom, seq) and not fullmatch(leadingEl, seq):
          return True
     return False

def isElectric(seq):
     electric = r'EL[0-9]{3}[0-9A-Z]{2}'
     if fullmatch(electric, seq):
          return True
     return False

#! Main Function
def analyze_license_plate(rawPlate):
     plate = preprocessPlate(rawPlate)

     if isValid(plate):
          if isStandard(plate):
               return "standard"
          elif isElectric(plate):
               return "electric"
          elif isCustom(plate):
               return "custom"
     return "invalid"
```

# Tests
The following cells contain public tests that you can use for basic validation of your solution. **Click on the validate button before submitting!**

```python
# This is a read-only cell used for automatic validation - do not edit, delete, or move!
import ipytest
ipytest.autoconfig(raise_on_error=True, addopts=["-p", "no:cacheprovider"])  # type: ignore
```

<!-- grade_id: cell-5be6ca123167ab1d, points: 2, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
# Short plates
short_plates = [
    "",
    " ",
    " 1 ",
    " 1A ",
    " 1A2 3 ",
    " 1A23 ",
    " 1A2 34 ",
    " 1A234 ",
    " 1A2 345 ",
    " 1A2345 ",
]
for plate in short_plates:
    try:
        analyze_license_plate(plate)
    except IndexError:
        print(f"TEST FAILED for plate {repr(plate)}")
        print("Try to analyze the function's input argument more carefully.")
        raise
    assert isinstance(analyze_license_plate(plate), str), "Your function does not return str object."
print("test OK")
```

<!-- grade_id: cell-4f592be6c2966f09, points: 2, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
plate_with_spaces = "   TRY PTHN1  "
assert analyze_license_plate(plate_with_spaces) == "custom", "You may have some problem with spaces."
print("test OK")
```

<!-- grade_id: cell-646329bb403be8c4, points: 2, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
plate_with_spaces = "2RY  HARD"
assert analyze_license_plate(plate_with_spaces) == "invalid", "You may have some problem with spaces."
print("test OK")
```

<!-- grade_id: cell-cb97fd8ba776eb28, points: 2, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
standard_plate = "1A10000"
assert analyze_license_plate(standard_plate) == "standard", "Check your criteria for standard plate."
print("test OK")
```

<!-- grade_id: cell-29abe45cd493fc93, points: 2, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
only_digits_plate = "1234567"
assert analyze_license_plate(only_digits_plate) == "custom", "Check your criteria for custom plate."
print("test OK")
```

<!-- grade_id: cell-b5627de0d6df87f0, points: 2, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
not_a_plate = "L1KEZPRO"
assert analyze_license_plate(not_a_plate) == "invalid", "Check your criteria for custom plate."
print("test OK")
```

<!-- grade_id: cell-c443f8e951424c73, points: 2, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
el_plate = "EL123AB"
assert analyze_license_plate(el_plate) == "electric", "Check your criteria for electric plate."
print("test OK")
```

<!-- grade_id: cell-173dbca087e7d864, points: 2, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
el_plate = "EL123AB0"
assert analyze_license_plate(el_plate) == "invalid", "Check your criteria for electric plate."
print("test OK")
```

<!-- grade_id: cell-91d40ea6836d9197, points: 5, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
invalid_plates = {
    "1S711-11": "invalid",
    "IAMTHE1!": "invalid",
    "L@VEZPR0": "invalid",
    "LOVEZPR0": "invalid",
    "1aA0987":  "invalid",
    "1AA":      "invalid",
    "4PZ 000":  "invalid",
    "3AČ0123":  "invalid",
    "1GA1234":  "invalid",
    "3AO 4567": "invalid",
    "2QQ 7891": "invalid",
    "4AW1111":  "invalid",
    "ABCDEFG":  "invalid",
    "123456":   "invalid",
    "123456789": "invalid",
}

for plate, result in invalid_plates.items():
    assert analyze_license_plate(plate) == result, "test case: " + plate
print("test OK")

```

<!-- grade_id: cell-ed6e34229ca9d376, points: 2, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
standard_plates = {
    "9BU 0000   ": "standard",
    "3A40123\n":   "standard",
    "\t2CH0987":   "standard",
    "1CH 2345":    "standard",
}

for test_case, result in standard_plates.items():
    assert analyze_license_plate(test_case) == result,  "test case: " + test_case
print("test OK")
```

<!-- grade_id: cell-6c5966bdb2b83529, points: 2, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
plates_with_spaces = {
    "   1AA0987   ": "standard",
    "1AA 0987":      "standard",
    "1AA  0987":     "invalid",
    "1A A0987":      "invalid",
    "ABCDE F1":      "invalid",
}

for test_case, result in plates_with_spaces.items():
    assert analyze_license_plate(test_case) == result, "test case: " + test_case
print("test OK")
```

<!-- grade_id: cell-cc8223b3783b5749, points: 5, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
custom_plates = {
    "L0V EZPR0": "custom",
    "FERRARI1":  "custom",
    "JAN H0NZA": "custom",
    "UV1D1ME":   "custom",
    "UVIDIME1":  "custom",
    "0PRAVIME":  "custom",
    "KABAT123":  "custom",
    "LACHTAN1":  "custom",
    "1234567":   "custom",
    "12345678":  "custom",
    "0AA0987   ": "custom",
}

for test_case, result in custom_plates.items():
    assert analyze_license_plate(test_case) == result, "test case: " + test_case
print("test OK")
```

<!-- grade_id: cell-54c915828ec333e6, points: 5, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
electric_plates = {
    "EL12345": "electric",
    "EL1 2345": "electric",
    "EL 12345": "invalid",
    "EL123ZZ": "electric",
    "EL123OO": "invalid",
    "EL121LE": "electric",
    "ELF ORIN": "invalid",
}

for test_case, result in electric_plates.items():
    assert analyze_license_plate(test_case) == result, "test case: " + test_case
print("test OK")
```

<!-- grade_id: cell-98b699c997d7a865, points: 5, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
region_plates = {
    "1A10001": "standard",
    "2A99999": "standard",
    "8A5A123": "standard",
    "9A9B999": "standard",
    "1D10001": "custom",
    "1F23456": "custom",
    "1I34567": "custom",
    "1N45678": "custom",
}

for test_case, expected in region_plates.items():
    assert analyze_license_plate(test_case) == expected, "region code test (custom): " + test_case

print("tests OK")
```

The following cells contain hidden tests that are evaluated during automatic grading after your submission. The tests check that your solution is sufficiently general.

<!-- grade_id: cell-9b8de79c66f5ecf1, points: 10, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
```

<!-- grade_id: cell-e1891c292a748924, points: 5, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
```

<!-- grade_id: cell-1fa910424b0a52e4, points: 5, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
```

<!-- grade_id: cell-366af49322244f38, points: 5, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
```

<!-- grade_id: cell-64cd4eaa47961188, points: 5, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
```

<!-- grade_id: cell-6a33ed35d53f8de8, points: 10, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
```

<!-- grade_id: cell-b6e679a72e2da410, points: 10, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
```

<!-- grade_id: cell-b3bd0fa4ce5bd905, points: 5, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
```

<!-- grade_id: cell-22ce285342036434, points: 5, grade: true, solution: false -->
```python
# Do not delete or edit this cell!
```